Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

**IMPORTANT: DO NOT COPY OR SPLIT CELLS.** If you do, you'll mess the autograder. If need more cells to work or test things out, create a new cell. You may add as many new cells as you need.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and group below:

```python
COURSE = "Unsupervised Learning 2021"
GROUP = "D8A"
NAME = "" # Match your GitHub Classroom ID
```

---

# Robust Principal Component Analysis

## Implementation (6 pts)

In this exercise, you will implement RPCA using the ADMM method reviewd in class. Then you will test it with synthetic data sets, and with handwritten digits.

Choose the step size for the ADMM algorithm as $\rho = \frac{nd}{4|M|_1}$.

Terminate the algorithm when $|M - L - S|_F \leq \delta |M|_F$, with $\delta=10^{-7}$.

This implementation is slow for large matrice, we could use a partial SVD to speed things up.

```python
import numpy as np
import matplotlib.pyplot as plt

%matplotlib inline
```

<!-- grade_id: ex-1, points: 0, grade: false, solution: true -->
```python
# Implement RPCA using the ADMM optimizing algorithm.

def soft_t(X, t):
    # Implement soft-thresholding
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
def SVT(L, t, k):
    # Implement Singular value thresholding
    
    # YOUR CODE HERE
    raise NotImplementedError()


def rpca(M, maxiter=10000):
    
    # You need to return the following variables correctly
    L = np.zeros(M.shape)
    S = np.zeros(M.shape)
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return L, S
```

```python
# Generate a synthetic low rank matrix of size nxn with n = 100
n = 100
r = int(0.05*n) # Rank
k = int(0.05*n**2) # Number of non-zero entries of S
X = np.random.normal(loc=0, scale=1/500, size=(n,r))
Y = np.random.normal(loc=0, scale=1/500, size=(n,r))
E = (-1)**np.random.randint(0, 2, size=(n, n))
S = np.zeros((n,n))
mask = np.array([True]*k + [False]*(n*n-k))
np.random.shuffle(mask)
mask = mask.reshape((n,n))
S[mask] = E[mask]

M = X @ Y.T + S
```

<!-- grade_id: ex-1-answer, points: 6, grade: true, solution: false -->
```python
# Check succesful recovery
Lr, Sr = rpca(M)
assert np.allclose(Lr, X@Y.T, atol=1e-6)
```

# Question 1 (2 pts)

Verify that the solution that minimizes

\begin{align}
\lambda|S|_1 + \frac{\rho}{2}\left|S - M + L_{k+1} - Y_k/\rho) \right|_F^2
\end{align}

with respect to $S$ is given by

$$
S_{k+1} = S_{\frac{\lambda}{\rho}}\left( M - L_{k+1} + \frac{Y_k}{\rho}  \right)
$$

<!-- grade_id: ans-q1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Question 1 (2 pts)

The right eigenvectors of the decomposition $\phi(X) = UDV^T$, i.e., the eigenvectors (loadings) in feature space, can be expanded in terms of the basis of observations,

$$
v_m = \sum_{j=1}^n \alpha_{jm}\phi(x_j)
$$

Show that the principal components for KPCA are given by

$$
z_{im} = v_m^T \phi(x_i) = \sum_{j=1}^n \alpha_{jm}\phi(x_j)^T \phi(x_i)
= \sum_{j=1}^n \alpha_{jm}K(x_i, x_j)
$$

with $\alpha_{jm} = u_{jm}/d_m$, assume a centered K matrix. This is, the vector of coefficients, $\alpha_m$ is equal to the $m$ eigenvector divided by the square of the eigenvalue, $\alpha_m = u_m/d_m$.

Also show that if the K matrix is not centered, we can still project any existing observation by centering K before, and any new observation $x_0$, by applying

$$
\vec{z}_{m} = A (I - M)\left(\vec{k}_0 - \frac{1}{N} K \vec{1}\right)
$$

where $A$ is the matrix of column vectors $\alpha_m$ given by $A=U D^{-1}$, and $k_0$ is a vector of kernel products $\phi(x_0)^T\phi(x_i)$ with all observations $x_i$.

<!-- grade_id: ans-q2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE